Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All). Check your output to make sure it all looks as you expected.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "DUCROCQ Romain"
```

---

You should have read all the readings and content from the lectures given on Brightspace.


In order to make it easy to grade your works, here are the notes you should read and follow carrefully.  We will take off points for not following them. We will grade based on assuming you have read and understood them.
    
1. **Using the Notebook to Show Your Work**: You must learn to write code in the notebook... It is a core tool for data science and will make it easier to develop and document your work if you become good at using it. Writing your code in another tool and pasting it into the notebook will probably not work well: Some of you forgot to include code elsewhere, or the spacing became messed up, or it didn't run as copied.  You must be sure your code cells execute because we will test them.  So learn how to run code in the notebook cells to double check your work.
2. **Read Directions Carefully**:  The instructions for your code are very important. If you don't follow the requirements, your application won't do as requested. Making it work correctly is part of learning to program.  If we worded something unclearly, send an e-mail to your professor.

The second set of issues are coding style:

1. **Indentation**: In Python, indentation matters and must be consistent. If you write your code in the Notebook, the tab key will indent properly.  If you use another editor and paste into the notebook, it might not be correctly indented (When you do write code in another editor, make sure you set your tabs to indent as 4 spaces, not as a tab character.)  You must make sure that your pasted code runs in the notebook or it will not get a good grade.
2. **Spacing**: Follow closely the spacing shown in the lessons.  There should not be a space  between a function name and the parentheses with the arguments. As a programmer, style is very important. If you work with programmers in the future, they sometimes have "lint" checkers to test your code for style and reject if it doesn't follow the approprate spacing and blank-line-rules. Think of it as a matter of politness for other people reading your code :-)
3. **Names of Variables**:  In Python, there's a culture of making everything readable. Don't use "x" and "y" as your variable names... use words like "pounds" and "kilograms."  It will be easier for colleagues (and yourself) to understand the code later.
4. **Error Messages**: Please use informative error messages that tell the user what they did wrong and what kind of input you expect.  Imagine you are designing the user experience!  Think about how to help your user. And remember **you** are the user when you debug!

We will take points off for issues of non-standard spacing, indentation, bad error messages, and bad variable names in the future.  This will continue for the entire course.

There are multiple ways to code all the answers.  Here are a few more code style tips:

1. If you do a calculation or a transformation, like "float(pounds)" -- do it once and save it as a variable, don't do it multiple times.  You should try not to have code that repeats itself too much.  If you repeat things, you can make mistakes like typos and it will be harder to find them.
2. A few of you had tests like "4 < test < 40" in your code, but didn't save that result (True) into any variable or put it in an IF statement.  It won't do anything relevant unless you save the result or use it as part of an IF test.
3. ``try``/``except`` should be used to catch errors. (In fancier, more formal Python, there is more careful error catching where the type of error is detected and handled. We're just doing the basic try/except right now.) Anytime you have a conversion or something that could result in an error, you should wrap it in try/except. Do not allow a user to run code that results in an un-handled error.

We will talk about this last point during in-class sessions. 
Once again, do not stay stuck without doing anything.
Use the "print" function to follow the elements and display your results inside your code.


---

# Homework Questions:

### Creation of a currency converter ###

Consider money you need to convert into a different currency. 
We will ask you to code a converter by asking the user to give an input value, and indicate the conversion he or she wants to make (from the initial currency to the final one).

The initial and final currencies will also be inputs from the user. 


#### First step####

Let's consider you as an user want to convert money from EUR to YEN (Renminbi).
Please take the currency value to be 1 EUR = 8.09647 YEN (on the 13/09/2020 - 1:00 pm / CET+2)

Write a code which: 
- Asks from the user an input considered as value money
This one will be considered to be in EUR 

- Then calculates the value in YEN 
- And displays and returns the result 
Here use both "print" and "return" functions

**Do not forgot to include tests / error messages / ...
in order to guide the user when using your program !**

All your error messages must start with "Error: " 


<!-- grade_id: cell-1b2e96b62467116f, points: 0, grade: false, solution: true -->
```python
# Remember you can comment your code using #

def currency_converter():
    ''' Ask for the user input with a prompt and compute the currency conversion from EUR to YEN. Print the value
    and return it from the function as well.'''
    
    prompt = 'Enter a value in EUR to be converted to YEN:'
    error_message = 'Error: your input should be a positive number'
    
    value = input(prompt)
    try:
        value = float(value)
    except:
        print(error_message)
        return None
    
    if value < 0:
        print(error_message)
        return None
    
    conversion = value * 8.09647
    print("Your input is equal to {output} stones".format(output=conversion))
    return conversion

```

To run your code in a cell, you should add a new cell and type:

```currency_converter()```

and execute that cell.  Provide the input and show it working below:

```python
# run this to show it working!
currency_converter()

```

Output:
```
Enter a value in EUR to be converted to YEN:i am a number !
Error: your input should be a positive number
```

The following code checks that your function performs correctly.  It is only verifying the return value is the correct number.  Testing functions for their return values is a standard practice in *Test-Driven Development*.  Test cases help you figure out where your code has errors, if you have lots of functions calling each other.

<!-- grade_id: test_converter1, points: 1, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

def test_converter():
    with mock.patch('builtins.input', return_value='4.5'):
        assert currency_converter() == 36.434115
    with mock.patch('builtins.input', return_value='200'):
        assert currency_converter() == 1619.294
            
test_converter()
```

Output:
```
Your input is equal to 36.434115 stones
Your input is equal to 1619.294 stones
```

<!-- grade_id: test_converter2, points: 2, grade: true, solution: false -->
```python
# testing more....

with mock.patch('builtins.input', return_value='cat'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert currency_converter() is None
        assert sys.stdout.getvalue() == "Error: your input should be a positive number\n"
```

<!-- grade_id: cell-c92e17942d63eb09, points: 1, grade: true, solution: false -->
```python
# testing even more....

with mock.patch('builtins.input', return_value='-3'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert currency_converter() is None
        assert sys.stdout.getvalue() == "Error: your input should be a positive number\n"
```

#### Second step####

Let's continue on the same topic.
Now consider multiple options of conversion from one currency to another one. 

This means by using what you've done preciously, add 1 input from the user to be: 
- the final currency for which you need to know the value
The initial currency of the value given will be stuck to EUR. 

For that purpose, you need to : 
- Store inputs into variables 
- Use inputs as elements for conditionals 
- Calculates conversion depending on the inputs
- Only use one output variable

For this exercise, you will consider the following currencies: YEN / USD / CAD / GBP / INR
The conversion rates are: 
- 1 EUR = 1.18468 USD
- 1 EUR = 0.925698 GBP 
- 1 EUR = 87.0200 INR
- 1 EUR = 1.56086 CAD

**When asking for inputs from the user, please stick to upper letters**
Please consider here inputs as arguments to the function currency_converter2()

<!-- grade_id: cell-acd5cd5cae046cf1, points: 0, grade: false, solution: true -->
```python
# Remember you can comment your code using #

def currency_converter2(value,currency):
    ''' Ask for the user input with a prompt and compute the currency conversion from EUR to YEN. Print the value
    and return it from the function as well.'''

    error_message = 'Error: your input should be a positive number'
    error_message2 = 'Error: select one of the given currency (YEN / USD / CAD / GBP / INR)'

    currencies = {"YEN": 8.09647, "USD": 1.18468, "CAD": 1.56086, "GBP": 0.925698, "INR": 87.0200}
    
    try:
        value = float(value)
    except:
        print(error_message)
        return None
    
    if value < 0:
        print(error_message)
        return None
    if not currency in [c for c in currencies]:
        print(error_message2)
        return None

    conversion = value * currencies[currency]
    print("Your input is equal to {output} ".format(output=conversion)+currency)
    return conversion

```

```python
# run this to show it is working!
# Give here content for the 2 variables value and currency
value = 50
currency = 'CAD'
currency_converter2(value, currency)
```

Output:
```
Your input is equal to 78.04299999999999 CAD
78.04299999999999
```

<!-- grade_id: test1_currency_converter2, points: 1, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

assert currency_converter2(3,'GBP') == 2.777094
assert currency_converter2(200,'USD') == 236.93599999999998
```

Output:
```
Your input is equal to 2.777094 GBP
Your input is equal to 236.93599999999998 USD
```

<!-- grade_id: test2_currency_converter2, points: 2, grade: true, solution: false -->
```python
# Some more tests
assert currency_converter2('cat','GBP') == None
assert currency_converter2(3,'Other') == None

```

Output:
```
Error: your input should be a positive number
Error: select one of the given currency (YEN / USD / CAD / GBP / INR)
```

<!-- grade_id: test3_currency_converter2, points: 1, grade: true, solution: false -->
```python
assert currency_converter2(0,'GBP') == 0

```

Output:
```
Your input is equal to 0.0 GBP
```

#### Third step####

During the last step, you will extend once again the program you've beeing programming. 

**Now consider not only one freezed input currency but multiple. 
This means the user need to specify in which currency the input is. 
He or she still has also to specify the output currency.** 

So here you need to: 
- Add one more input from the user specifying the input or initial currency.

Once again store this input into a variable and use it for calculating the output value. 

**Take time before coding to consider all the options of conversion, from one currency to another one. And try a way to calculate each conversion even if it is not given here in the instructions.** 

Concerning the conversion rates, use the ones indicate in 2nd step. 



<!-- grade_id: cell-e6ae9b5ba8031bea, points: 0, grade: false, solution: true -->
```python
# Remember you can comment your code using #
# You can see that here we switch from input asking to arguments giving directly to the function

def currency_converter3(value, currency1, currency2):
    ''' From a value given by the user the function computes the currency conversion from one currency to another one. Print the value
    and return it from the function as well.'''
    
    error_message = 'Error: your input should be a positive number'
    error_message2 = 'Error: select one of the given currency (EUR,YEN,USD,CAD,INR,GBP)'
    
    currencies = {"EUR": 1.0, "YEN": 8.09647, "USD": 1.18468, "CAD": 1.56086, "GBP": 0.925698, "INR": 87.0200}
    
    try:
        value = float(value)
    except:
        print(error_message)
        return None
    
    if value < 0:
        print(error_message)
        return None
    if not (currency1 in [c for c in currencies] and currency2 in [c for c in currencies]):
        print(error_message2)
        return None

    conversion = value * currencies[currency2] / currencies[currency1]

    print(conversion)
    return conversion

```

<!-- grade_id: cell-c788fe9305e7f4f7, points: 2, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

assert currency_converter3(200,'INR','USD') == 2.722776373247529
assert currency_converter3(30,'CAD','GBP') == 17.792076163140834

```

Output:
```
2.722776373247529
17.792076163140834
```

<!-- grade_id: cell-16120a64ef493730, points: 2, grade: true, solution: false -->
```python
# Some more tests
assert currency_converter3(3,'cat','GBP') == None
assert currency_converter3(3,'EUR','Other') == None

```

Output:
```
Error: select one of the given currency (EUR,YEN,USD,CAD,INR,GBP)
Error: select one of the given currency (EUR,YEN,USD,CAD,INR,GBP)
```

## You're Done!  Upload this file to your dropbox (double check your name is at the top, and the Notebook title hasn't been changed). 